# Add function

Create a function that adds two whole numbers, ValueError should be raised on invalid params

<!-- grade_id: answer-cell, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    ### BEGIN SOLUTION CODE
    if not isinstance(a, int) or not isinstance(b, int):
        raise ValueError
    return a + b
    ### END SOLUTION CODE
    raise NotImplemented
```

```python
def throwError():
    raise ValueError
```

<!-- grade_id: shown-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 0, Expected Result: pass
assert True, "Sample hint message 1"
```

<!-- grade_id: shown-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 1, Expected Result: fail
assert 0 > 1, "Sample hint message 2"
```

<!-- grade_id: shown-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 1, Expected Result: fail
assert throwError(), "Error thrown"
```

<!-- grade_id: hidden-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: pass

### BEGIN HIDDEN TESTS
assert True, "Sample hint message 3"
### END HIDDEN TESTS
```

<!-- grade_id: hidden-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
### BEGIN HIDDEN TESTS
add(4, "string")
assert "False" == "True", "Sample hint message 4"
### END HIDDEN TESTS
```

<!-- grade_id: hidden-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
add("string", 4)
### BEGIN HIDDEN TESTS
assert throwError()
### END HIDDEN TESTS
```

<!-- grade_id: shown-fail-hidden-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert len("string") == 10, "Sample hint message 5"
### BEGIN HIDDEN TESTS
assert True, "Sample hint message 6"
### END HIDDEN TESTS
```

<!-- grade_id: shown-pass-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert "True".isalpha(), "Sample hint message 7"
### BEGIN HIDDEN TESTS
assert False, "Sample hint message 8"
### END HIDDEN TESTS
```

<!-- grade_id: hidden-pass-shown-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
### BEGIN HIDDEN TESTS
assert True, "Sample hint message 9"
### END HIDDEN TESTS
assert False
```

<!-- grade_id: hidden-fail-shown-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
### BEGIN HIDDEN TESTS
assert False, "Sample hint message 11"
### END HIDDEN TESTS
assert True
```

<!-- grade_id: shown-fail-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert False, "Sample hint message 13"
### BEGIN HIDDEN TESTS
assert False, "Sample hint message 14"
### END HIDDEN TESTS
```

```python
!cat *.ipynb
```