# Add function

Create a function that adds two whole numbers, ValueError should be raised on invalid params

<!-- grade_id: answer-cell, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    # your code here
    
    raise NotImplemented
```

```python
def throwError():
    raise ValueError
```

<!-- grade_id: shown-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 0, Expected Result: pass
assert True
```

<!-- grade_id: shown-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 3, Expected Result: fail
assert False
# Hint: test hint for R
 #Hint: Weird spacing
    #Hint: Tab before comment
```

<!-- grade_id: shown-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 1, Expected Result: fail
assert throwError()
# Hint: this test should throw an error
```

<!-- grade_id: hidden-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: pass

```

<!-- grade_id: hidden-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
#Hint: hint 2
```

<!-- grade_id: hidden-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
```

<!-- grade_id: shown-fail-hidden-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
assert False
```

<!-- grade_id: shown-pass-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 3, Expected Result: fail
assert True
#Hint: hint 1
# Hint: hint2
#  Hint: hint 3
```

<!-- grade_id: hidden-pass-shown-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
assert False
```

<!-- grade_id: hidden-fail-shown-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
assert True
```

<!-- grade_id: shown-fail-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
assert False
```

```python
!cat *.ipynb
```