# Add function

Create a function that adds two whole numbers, ValueError should be raised on invalid params

<!-- grade_id: answer-cell, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    # your code here
    
    raise NotImplemented
```

```python
def throwError():
    raise ValueError
```

<!-- grade_id: shown-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 0, Expected Result: pass
assert True, "Sample hint message 1"
```

<!-- grade_id: shown-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 1, Expected Result: fail
assert 0 > 1, "Sample hint message 2"
```

<!-- grade_id: shown-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: False, Expected Number of Hints: 1, Expected Result: fail
assert throwError(), "Error thrown"
```

<!-- grade_id: hidden-pass, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: pass

```

<!-- grade_id: hidden-fail, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
```

<!-- grade_id: hidden-error, points: 1, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 0, Expected Result: fail
add("string", 4)
```

<!-- grade_id: shown-fail-hidden-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert len("string") == 10, "Sample hint message 5"
```

<!-- grade_id: shown-pass-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert "True".isalpha(), "Sample hint message 7"
```

<!-- grade_id: hidden-pass-shown-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
assert False
```

<!-- grade_id: hidden-fail-shown-pass, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 1, Expected Result: fail
assert True
```

<!-- grade_id: shown-fail-hidden-fail, points: 2, grade: true, solution: false -->
```python
# Hidden Tests: True, Expected Number of Hints: 2, Expected Result: fail
assert False, "Sample hint message 13"
```

```python
!cat *.ipynb
```