# Add function

Create a function that adds two whole numbers, ValueError should be raised on invalid params

<!-- grade_id: answer-cell, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    return a + b
```

<!-- grade_id: cell-bcf433e35f021efe, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: shown-pass, points: 1, grade: true, solution: false -->
```python
assert True
```

<!-- grade_id: shown-fail, points: 1, grade: true, solution: false -->
```python
assert False
```

<!-- grade_id: shown-error, points: 1, grade: true, solution: false -->
```python
assert throwError()
```

<!-- grade_id: hidden-pass, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: hidden-fail, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: hidden-error, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: shown-fail-hidden-pass, points: 2, grade: true, solution: false -->
```python
assert False
```

<!-- grade_id: shown-pass-hidden-fail, points: 2, grade: true, solution: false -->
```python
assert True
```

<!-- grade_id: hidden-pass-shown-fail, points: 2, grade: true, solution: false -->
```python
assert False
```

<!-- grade_id: hidden-fail-shown-pass, points: 2, grade: true, solution: false -->
```python
assert True
```

<!-- grade_id: shown-fail-hidden-fail, points: 2, grade: true, solution: false -->
```python
assert False
```